### Sahil Shivaji Patil
#### 21510060
#### T-3
#### Assignment 5

```python
pip install nose
```

Output:
```
Collecting nose
  Using cached nose-1.3.7-py3-none-any.whl (154 kB)
Installing collected packages: nose
Successfully installed nose-1.3.7
Note: you may need to restart the kernel to use updated packages.
```

```python
#Perform logistic regression on diabetes dataset
```

<!-- grade_id: cell-65901a02d7b3d454, points: 0, grade: false, solution: true -->
```python
#Load diabetes dataset here and return shape of dataset
from cz_utility import dataset
# diabetes dataset
df = dataset.load('63ea33b29ac5f900139b8d07')
def returnShape(df):
    # your code here
    return df.shape
    raise NotImplementedError
df.columns
```

Output:
```
Index(['Pregnancies', 'Glucose', 'BloodPressure', 'SkinThickness', 'Insulin',
       'BMI', 'DiabetesPedigreeFunction', 'Age', 'Outcome'],
      dtype='object')
```

```python
df
```

Output:
```
Pregnancies  Glucose  BloodPressure  SkinThickness  Insulin   BMI  \
0              6      148             72             35        0  33.6   
1              1       85             66             29        0  26.6   
2              8      183             64              0        0  23.3   
3              1       89             66             23       94  28.1   
4              0      137             40             35      168  43.1   
..           ...      ...            ...            ...      ...   ...   
763           10      101             76             48      180  32.9   
764            2      122             70             27        0  36.8   
765            5      121             72             23      112  26.2   
766            1      126             60              0        0  30.1   
767            1       93             70             31        0  30.4   

     DiabetesPedigreeFunction  Age  Outcome  
0                       0.627   50        1  
1                       
```

<!-- grade_id: cell-9d4cd5b17d896c66, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(returnShape(df), (768, 9))
```

<!-- grade_id: cell-97b7b8d9c24aa579, points: 0, grade: false, solution: true -->
```python
#separate input and output from dataset and remove unwanted data columns
#take x for input and y for output
def sepAndRemove(df):
    # your code here
    X=df.drop(columns=['Outcome','SkinThickness'])
    y=df['Outcome']
    return X,y
    raise NotImplementedError
```

<!-- grade_id: cell-f7e39f11fcd55f39, points: 2, grade: true, solution: false -->
```python
trainAndTestData = sepAndRemove(df)
assert_equal(trainAndTestData[0].shape, (768, 7))
assert_equal(trainAndTestData[1].shape, (768,))
```

<!-- grade_id: cell-05c4eb3e485c557a, points: 0, grade: false, solution: true -->
```python
#separate training and testing dataset with test_size = 0.3
from sklearn.model_selection import train_test_split
def trainTestSplit(trainAndTestData):
    # your code here
    x_train,x_test,y_train,y_test = train_test_split(trainAndTestData[0],trainAndTestData[1],test_size=0.3,random_state=42)
    return x_train,x_test,y_train,y_test
    raise NotImplementedError
```

<!-- grade_id: cell-a7a6e742dbef0d20, points: 2, grade: true, solution: false -->
```python
trainAndTestData=trainTestSplit(sepAndRemove(df))
assert_equal(trainAndTestData[0].shape,(537, 7))
assert_equal(trainAndTestData[1].shape,(231, 7))
assert_equal(trainAndTestData[2].shape,(537,))
assert_equal(trainAndTestData[3].shape,(231,))
```

<!-- grade_id: cell-6e973aeb54f8a469, points: 0, grade: false, solution: true -->
```python
#use MinMaxScaler to scale the data to a fixed range
from sklearn.preprocessing import MinMaxScaler
def scalar(trainAndTestData):
    # your code here
    scaler=MinMaxScaler()
    x_trainscaled=scaler.fit_transform(trainAndTestData[0])
    x_testscaled=scaler.transform(trainAndTestData[1])
    return x_trainscaled,x_testscaled,trainAndTestData[2],trainAndTestData[3]
    raise NotImplementedError
```

<!-- grade_id: cell-0bb62bf101293688, points: 2, grade: true, solution: false -->
```python
trainAndTestData=scalar(trainTestSplit(sepAndRemove(df)))
assert_equal(trainAndTestData[0].shape, (537, 7))
assert_equal(trainAndTestData[1].shape, (231, 7))

```

```python
trainAndTestData[0].shape,trainAndTestData[1].shape
```

Output:
```
((537, 7), (231, 7))
```

<!-- grade_id: cell-a9178e0adf01003e, points: 0, grade: false, solution: true -->
```python
#perform logistic regression here and return accuracy between predected output for testx and testy
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score
def logisticReg(trainAndTestData, trainy, testy):
    # your code here
    log_regressor=LogisticRegression()
    log_regressor.fit(trainAndTestData[0],trainy)
    predy=log_regressor.predict(trainAndTestData[1])
    return accuracy_score(testy,predy)
    raise NotImplementedError
```

<!-- grade_id: cell-b26830e8ae1d3a06, points: 2, grade: true, solution: false -->
```python
#there might be difference between student answer and faculty answer
trainAndTestData1 = trainTestSplit(sepAndRemove(df))
accuracyScore = logisticReg(scalar(trainTestSplit(sepAndRemove(df))), trainAndTestData1[2], trainAndTestData1[3])
assert_equal(accuracyScore>0.60, True)
```

```python
accuracyScore
```

Output:
```
0.7532467532467533
```

```python

```